# 1. feladatsor - Dekorátorok

## Oszthatóság számolása

Készíts egy Python dekorátort `timeit_decorator` néven ami képes lemérni es printelni egy függvény időkomplexitását, azaz azt az időt ami a futtatásához kellett. 

PL:

```
@timeit_decorator
def comprehension():

    return [x**2 for x in range(3000000)]
```

A kiírás elvárt formája:

`FUNCTION_NAME took TIME`


`comprehension took 767.8432464599609`


<!-- grade_id: cell-84721ecb4b633c21, points: 0, grade: false, solution: true -->
```python
import time

# YOUR CODE HERE
raise NotImplementedError()

@timeit_decorator
def comprehension():
    return [x**2 for x in range(3000000)]

@timeit_decorator
def generator():
    return (x**2 for x in range(3000000))

@timeit_decorator
def iter_over(iter):
    for _ in iter:
        pass

    
alist = comprehension()
iter_over(alist)
agenerator = generator()
iter_over(agenerator)
```

<!-- grade_id: cell-0bede469175bd15d, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal(len(solution), 4)
assert_true("comprehension took" in solution[0])
assert_true("iter_over took" in solution[1])
assert_true("generator took" in solution[2])
assert_true("iter_over took" in solution[3])
```

## Anonimizáló dekorátor
Készíts egy dekorátort `anonimize` néven, ami egy függvény visszatérési értékét anonimizálja, azaz olvashatatlanná teszi, ha a visszatérési érték `str` típusú.
Tegye ezt úgy, hogy minden karaktert `#`-é alakít.
Pl:
    
`almafa -> ######`




<!-- grade_id: cell-5c227f0b8eec859e, points: 0, grade: false, solution: true -->
```python
import random
# YOUR CODE HERE
raise NotImplementedError()

@anonimize
def get_result():
    data = ["alma", "kortefa", "barack", "bananfa"]
    return random.choice(data)


print(get_result())
```

<!-- grade_id: cell-6658e55c4a75eb0e, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true('#' in get_result())
assert_true('#' in get_result())
assert_true('#' in get_result())
assert_true('#' in get_result())
assert_true('#' in get_result())
```

## Logging dekorátor

Írj egy olyan dekorátort `logged` néven, ami egy változó hosszú paraméter listával rendelkező függvényt be tud csomagolni és ki tudja írni a függvény összes paraméterét, nevét, valamint azt, hogy hány paramétere van.

Kiírási minta:

```python
hello(2,4,6)
=>
you called hello((2, 4, 6))
result was 6
```

<!-- grade_id: cell-ab64cbd1a8753849, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

@logged
def hello(*args):
    return 3 + len(args)

hello(2,4,6)
```

<!-- grade_id: cell-be785a9ad12141b5, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal(solution[0], "you called hello((2, 4, 6))")
assert_equal(solution[1], "result was 6")

```

## 2. feladatsor - Haladó objektum orientált programozás

Készíts egy Python programot ami az alábbi osztályokat és interface-eket valósítja meg:

```ascii
    ┌─────────────────────────┐            ┌─────────────────────────┐
    │       <<abc.ABC>>       │            │                         │
    │          Bird           │            │         Parrot          │
    ├─────────────────────────┤            ├─────────────────────────┤
    │                         │◄───────────┤                         │
    ├─────────────────────────┤            ├─────────────────────────┤
    │         *fly()          │            │         *fly()          │
    └─────────────────────────┘            └─────────────────────────┘
    ┌─────────────────────────┐            ┌─────────────────────────┐
    │       <<abc.ABC>>       │            │                         │
    │          Hero           │            │        SuperMan         │
    ├─────────────────────────┤            ├─────────────────────────┤
    │                         │◄───────────┤                         │
    ├─────────────────────────┤            ├─────────────────────────┤
    │         *fly()          │            │         *fly()          │
    ├─────────────────────────┤            ├─────────────────────────┤
    │       *use_power()      │            │       *use_power()      │
    └─────────────────────────┘            └─────────────────────────┘
```

Ügyelj rá, hogy a formális interface-t használj.

A példa futása az alábbi kimenetet kell adnia:
```text
slow flyer
superpower
False
False
```

<!-- grade_id: cell-959a7afa3ad2a3fa, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()


p = Parrot()

p.fly()

super_man = SuperMan()

super_man.use_power()

print(isinstance(p, Hero))

print(isinstance(super_man, Bird))
```

<!-- grade_id: cell-d63bdf5f3da38857, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_false(isinstance(p, Hero))
assert_false(isinstance(super_man, Bird))
assert_true(isinstance(p, Bird))
assert_true(isinstance(super_man, Hero))
assert_true(hasattr(super_man, 'use_power'))
assert_true(hasattr(super_man, 'fly'))
assert_true(hasattr(p, 'fly'))
assert_false(hasattr(p, 'use_power'))
```

## 3. feladatsor - többszörös öröklődés
Készíts el egy Python programot ami az alábbi objektum hierarchiát és működést valósítja meg:

```ascii
                         +---------------------------+
                         |          Address          |    +---------------------------+
                         +---------------------------+    |          Person           |
                         |          street           |    +---------------------------+
                         |           city            |    |                           |
                         +---------------------------+    +---------------------------+
                         |__init__(street, city)     |    |__init__(name, email)      |
                         |show()                     |    |show()                     |
                         +-------------+-------------+    +-------------------+-------+
                                       |                                      |
                                       |                                      |
                                       |                                      |
                                       |                                      |
                                       |                                      |
+---------------------------+          +>+-----------------------------+<-----+
|         Notebook          |            |         Contact             |
+------------------------+--+  1   0..1  +-----------------------------+
|    people:dict(Contact)|+++----------->|                             |
+------------------------+--+            +-----------------------------+
|add(name, email, street,   |            |__init__(name, email, street,|
|    city)                  |            |         city)               |
|show(name)                 |            |show()                       |
+---------------------------+            +-----------------------------+
```

A többes öröklődés a Contact, Address és Person objektumok között áll fenn. A teszteléshez egy Notebook osztályt készíts ami használja a Contact osztályt. Tárold a ”people” nevű példány változóban a Contact-otkat. Ez legyen egy dict. A show metódus a Notebookon legyen képes megjeleníteni az egyes Contact-okat név alapján.

Elvárt kimenet:

```text
Alice <al@kth.se>
Lv 24
Sthlm

Unknown Carol
```

<!-- grade_id: cell-8212d83c0e0ccc47, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
notes = Notebook()
notes.add('Alice', '<al@kth.se>', 'Lv 24', 'Sthlm')
notes.add('Bob', '<bb@kth.se>', 'Rtb 35', 'Sthlm')

notes.show('Alice')
notes.show('Carol')
```

<!-- grade_id: cell-ab973b6234b288bc, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true(type(notes.people) == dict)
assert_true(type(notes.people['Alice']) == Contact)
assert_true(isinstance(notes.people['Alice'],Person))
assert_true(isinstance(notes.people['Alice'],Address))
```