# 1. feladatsor - Asszinkron programozás
## Chat program szerver és kliens

Készíts egy Python chat alkalmazást asyncio segitségével. Legyen külon egy kliens és egy sever modul. Tetszőleges üzenettel lehessen bontani a kapcsolatot. A chat üzenetekete a szerver és a kliens is a saját szemszögéből jelenítse meg.

Ezeket a feladatokat nem tudod ebben a környezetben futtatni. A kódot külön fájlokban írd be saját környezetben, és vesd össze az itt található megoldással.

<!-- grade_id: cell-9897d952a40b410e, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-82217a505b64d24e, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# 2. feladatsor - RXPY reaktív programozás
## Tőzsdei vevő jelző

Készíts Python alkalmazást, ami az RxPy modul használatával valósít meg egy egyszerű tőzsdei vevő jelző rendszert.

Legyen egy listád a különböző részvényekről:
```python
stocks = [
{ 'TCKR' : 'APPL', 'PRICE': 200},
{ 'TCKR' : 'GOOG', 'PRICE': 90},
{ 'TCKR' : 'TSLA', 'PRICE': 120},
{ 'TCKR' : 'MSFT', 'PRICE': 150},
{ 'TCKR' : 'INTL', 'PRICE': 70},
]
```

Készíts ebből egy observable forrást. A forrás küldjön instrukciókat a feliratkozottaknak (observer függvény) de csak akkor ha a részvény árfolyam éppen 100 USD felett van.

Az observer függvényed és így a programod is ilyen kimenetet kell adjon:
```text
Received Instruction to buy APPL

Received Instruction to buy TSLA

Received Instruction to buy MSFT

All Buy Instructions have been received

```

<!-- grade_id: cell-9fa7d20a5c2a5066, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bd3e30a89ae3fbb3, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal(solution, ['Received instruction to buy APPL', 'Received instruction to buy TSLA', 'Received instruction to buy MSFT', 'All buy instructions have been received'])
```