# 1. feladatsor - listák

## Adott az alábbi lista:
 ​
`lista = [-25, 12, -54, 8, 77, 98, -29, 35, 3, 71]`

Írj ciklust, amely meghatározza, melyik a legnagyobb szám a listában, és írd ki ezt a számot! (Melyik elemeket kell a ciklusnak vizsgálnia? Vigyázat: nem kell az összeset! Az egy elemű listának annak egyetlen eleme a maximuma.)
 ​

Módosítsd úgy a programot, hogy a ciklusba belépés előtt, és a ciklustörzs végén is kiírod mindig az épp megvizsgált elemet, és a maximum változó értékét! Vizsgáld meg az eredményt, figyeld meg, a maximum változó hogyan változik!

<!-- grade_id: cell-a3e8d85db2cd3829, points: 0, grade: false, solution: true -->
```python
lista = [-25, 12, -54, 8, 77, 98, -29, 35, 3, 71]
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d8bf68236e9f6fa6, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true(int(solution[0]) == 98)
```

## Műveletek listákkal

Írj programot, amely tartalmaz egy tíz elemű listát, az általad megadott kezdeti értékekkel inicializálva! (Tehát nem kell a programnak egyesével beolvasnia azokat a billentyűzetről.) Írd ki ezt a listát, legyen a neve `lista`!

Kiíratás:
* `A lista: 25 69 54 8 77 6 29 10 3 98`

Írd át úgy a programot, hogy megjelenjen a listaelemek előtt a listaindex is!

Kiíratás:
* `A lista: [0]=25 [1]=69 [2]=54 [3]=8 [4]=77 [5]=6 [6]=29 [7]=10 [8]=3 [9]=98`


Írj programrészt, amely megmondja, melyik a legkisebb szám a listából! Írd ki ezt a számot! (Próbáld ki a ogramod úgy is, hogy a legkisebb szám a lista legelején és legvégén van!)


Kiíratás:
* `A legkisebb szám: 3`


Alakítsd át a minimumkeresést úgy, hogy ne csak a legkisebb szám értékét, hanem annak helyét, azaz listabeli indexét is meg tudd mondani! Ehhez szükséged lesz egy minhely változóra, amely azt fogja megjegyezni, hányadik indexen volt a legkisebb szám.


Az igazán jó megoldás az, ha továbbra is egy ciklusod lesz, amelyik ezt a keresést végzi. Tehát nem úgy kell működnie, hogy előbb megjegyzi a legkisebb számot (3), utána pedig újból végigszalad a listán, hogy rájöjjön, hol is volt az (8-as hely). A helyet már a keresés közben is meg lehet jegyezni. A futási eredmény legyen ilyen:



Kiíratás:
* `A legkisebb szám indexe: 8`


Végül pedig írd úgy ki a listát, hogy a legkisebb elem mellé egy jelölést teszel:


Kiíratás:
* `Jelölve: 25 69 54 8 77 6 29 10 3[MIN] 98`


<!-- grade_id: cell-84721ecb4b633c21, points: 0, grade: false, solution: true -->
```python
# Itt ha mar az elozo feladatban megirtuk az is_szokoev fuggvenyt, akkor azt tudjuk hivatkozni, nem kell idemasolni
lista = [25, 69, 54, 8, 77, 6, 29, 10, 3, 98]

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0bede469175bd15d, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true(len(solution) == 5)
assert_equal(solution[0].strip(), "A lista: 25 69 54 8 77 6 29 10 3 98")
assert_equal(solution[1].strip(), "A lista: [0]=25 [1]=69 [2]=54 [3]=8 [4]=77 [5]=6 [6]=29 [7]=10 [8]=3 [9]=98")
assert_equal(solution[2].strip(), "A legkisebb szám: 3")
assert_equal(solution[3].strip(), "A legkisebb szám indexe: 8")
assert_equal(solution[4].strip(), "Jelölve: 25 69 54 8 77 6 29 10 3[MIN] 98")
```

## Lista megfordítás

A listák .reverse() függvénye megfordítja a listát, amelyre meghívják:
```python
szamok = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]
szamok.reverse()
print(szamok)
```
`[10, 9, 8, 7, 6, 5, 4, 3, 2, 1]`


**Készíts listákat megfordító algoritmusokat, amelyek a következőképpen működnek:**

Adott egy lista az eredeti nevű változóban. Járd be ezt visszafelé (végétől az elejéig), és fűzd hozzá a forditott nevű listához a benne tárolt adatokat!


Kiíratás:
* `A forditott lista [10, 9, 8, 7, 6, 5, 4, 3, 2, 1]`


Adott egy újabb lista az eredeti nevű változóban. Minden végéről kivett elemet fűzz hozzá egy új lista végéhez!


Kiíratás:
* `A forditott lista [10, 9, 8, 7, 6, 5, 4, 3, 2, 1]`


Fordítsd meg a következő listát helyben, azaz ne használj másik változót a megoldás során! Cseréld meg az első és az utolsó, a második és az utolsó előtti, ... elemét, így kapva meg a fordított listát!

`lista = [9, 8, 7, 6, 5, 4, 3, 2, 1]`

Kiíratás:
* `A helyben megforditott lista [9, 8, 7, 6, 5, 4, 3, 2, 1]`

<!-- grade_id: cell-3d41946d709432c4, points: 0, grade: false, solution: true -->
```python
lista = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e3e4a0a8c8e2be2f, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal(solution[0].strip(), "A forditott lista [10, 9, 8, 7, 6, 5, 4, 3, 2, 1]")
assert_equal(solution[1].strip(), "A forditott lista [10, 9, 8, 7, 6, 5, 4, 3, 2, 1]")
assert_equal(solution[2].strip(), "A helyben megforditott lista [9, 8, 7, 6, 5, 4, 3, 2, 1]")
```