# 1. feladatsor - set
## Lista elemek szűrése

Adott az alábbi lista

`sample_list = [87, 45, 41, 65, 94, 41, 99, 94]`

Távolítsuk el a duplikált elemeket belőle és rendezzük növekvő sorrendbe. Írjuk ki az elemeket.

Kiírás:
* `44 52 53 54 87`

Keressük meg a minimum és maximum értéket. Ezeket írjuk ki a képernyőre.

Kiírás:
* `44`
* `87`

<!-- grade_id: cell-84721ecb4b633c21, points: 0, grade: false, solution: true -->
```python
sample_list = [87, 52, 44, 53, 54, 87, 52, 53]
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0bede469175bd15d, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true(len(solution) == 3)
assert_equal(solution[0].strip(), "44 52 53 54 87")
assert_equal(solution[1].strip(), "44")
assert_equal(solution[2].strip(), "87")
```

## Párok

Adott két lista aminek azonos a mérete. Írd ki az összes olyan párt amit a két lista elemeiből lehet alkotni. Csak az egyedi párokat tartsd meg. Az eredményt tárold a `result` változóban!


Első lista `[2, 3, 4, 5, 6, 7, 8]`

Második lista `[4, 9, 16, 25, 36, 49, 64]`

Eredmény `{(6, 36), (8, 64), (4, 16), (5, 25), (3, 9), (7, 49), (2, 4)}`

<!-- grade_id: cell-1bc5a980cd583d68, points: 0, grade: false, solution: true -->
```python
first_list = [2, 3, 4, 5, 6, 7, 8]
second_list = [4, 9, 16, 25, 36, 49, 64]
result = None
# YOUR CODE HERE
raise NotImplementedError()
print(result)
```

<!-- grade_id: cell-4d96a53f5f5e00fb, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal(result, {(7, 49), (2, 4), (4, 16), (8, 64), (6, 36), (3, 9), (5, 25)})
```