# Jupyter notebook bevezető
Ez a munkafüzet egy interaktív felület, ahol a jelölt rubrikákba kódot írva le lehet futtatni azokat. Az adott kódrészletet futtatni a felső menüben található gombokkal lehet (Run).

A megoldásaidat az adott cellába írd, töröld a
```python
    # YOUR CODE HERE
    raise NotImplementedError()
```

Megoldás után az adott és a következő cellát futtatsd (ezek az automatikus tesztek). Ha nem ír semmit, akkor jó a megoldás, ha hibát kapsz akkor javítsd az megoldást és futtasd mindkét cellát újra.

Ha a notebook végéig minden teszt hiba nélkül lefut, akkor kész a feladatsor.

# Hibaüzenetek, hibakezelés

Néhány tipikus hibaüzenet, ha esetleg elakadnál. Nem kell most elolvasnod, csak ha hibát kapsz és nem érted, akkor itt ez esetleg segíthet.

* `NotImplementedError` -> Nem törölted ki a # YOUR CODE HERE részt
* `AssertionError` -> Nem jó a feladat megoldása, a teszt hibára futott
* `NameError` -> Rosszul használod a Python parancsait, nézd át a kódot


## Példák

```python
---------------------------------------------------------------------------
NotImplementedError                       Traceback (most recent call last)
<ipython-input-2-6d4db73a6e75> in <module>
      2 
      3     # YOUR CODE HERE
----> 4     raise NotImplementedError()
      5 
      6 print(solution[0])

NotImplementedError: 
```

--> Nem törölted ki a kért részt, ki kell törölni és a helyére gépelni a megoldásodat.

```python
---------------------------------------------------------------------------
AssertionError                            Traceback (most recent call last)
<ipython-input-4-28814bafb7f0> in <module>
      3 solution = get_ipython().run_line_magic('solution', '')
      4 assert_true(len(solution))
----> 5 assert_equal(float(solution[0]),  31200)

/opt/conda/lib/python3.9/unittest/case.py in assertEqual(self, first, second, msg)
    829         """
    830         assertion_func = self._getAssertEqualityFunc(first, second)
--> 831         assertion_func(first, second, msg=msg)
    832 
    833     def assertNotEqual(self, first, second, msg=None):

/opt/conda/lib/python3.9/unittest/case.py in _baseAssertEqual(self, first, second, msg)
    822             standardMsg = '%s != %s' % _common_shorten_repr(first, second)
    823             msg = self._formatMessage(msg, standardMsg)
--> 824             raise self.failureException(msg)
    825 
    826     def assertEqual(self, first, second, msg=None):

AssertionError: 312200.0 != 31200
```

--> Az automatikus teszt hibát talált a megoldásodban.


```python
---------------------------------------------------------------------------
NameError                                 Traceback (most recent call last)
<ipython-input-4-ba4712915cd0> in <module>
      1 from nose.tools import assert_equal, assert_true
      2 
----> 3 assert_true(len(solution))
      4 assert_equal(float(solution[0]),  31200)

NameError: name 'solution' is not defined
```
--> Nem futtattad a megoldás celláját, csak a tesztét. Mindenképpen először a megoldás cellát kell futtatni, utána a tesztet. A bal oldali `In []` feliratot kell nézni, ha itt a [] ben van szám, akkor futtatva volt.

```python
---------------------------------------------------------------------------
NameError                                 Traceback (most recent call last)
<ipython-input-10-f48fdeef71f0> in <module>
      2 
      3     
----> 4     prit(90 / 15 * 5200)
      5     
      6 

NameError: name 'prit' is not defined
```

--> Rosszul használsz valamilyen Python kulcsszót, itt pl el van írva a `print` helyett `prit`

```python
---------------------------------------------------------------------------
TypeError                                 Traceback (most recent call last)
<ipython-input-12-1e94ef827a25> in <module>
      7     eletkor = input("Add meg az életkorod: ")
      8     ital = input("Milyen italt kérsz?: ")
----> 9     if eletkor < 18 and ital == "sör":
     10         print("Sajnos nem adhatok.")
     11         exit(1)

TypeError: '<' not supported between instances of 'str' and 'int'
```

--> Rossz típusokat használsz, itt a hiba pl, hogy az `input` eredményét először számmá kell konvertálni `int(input("valami"))`

# 1. feladatsor - egyszerű utasítások 
## Szoba kifestés
A 90m<sup>2</sup>-es szobádat akarod kfiesteni, egy vödör festék 15<sup>2</sup> kifestésére elég és 5200 Ft-ba kerül.

**Mennyibe fog kerülni a teljes szoba kifestése?**

<!-- grade_id: 1_feladat, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: 1_feladat_test, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true(len(solution))
assert_equal(float(solution[0]),  31200)
```

## Festő bére
A festő egy m<sup>2</sup> -t 12 perc alatt fest le, de az egész szobát kétszer festi át. Ha nettó 1800 forintos órabérrel dolgozik, akkor mennyi lesz a nettó végösszeg?

Mennyi lesz 27% áfával a végösszeg?

<!-- grade_id: 2_feladat, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: 2_feladat_test, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true(len(solution) == 2)
assert_equal(float(solution[0]),  13500)
assert_equal(float(solution[1]),  17145)
```

## Nagymama befőttjei
Nagymama 30kg paradicsomot akar befőzni. Kilónként 0.8 liter paradicsompüré lesz a főzés során, és ezeket 6 decis üvegekbe tölti.

**Hány üveget tud megtölteni teljesen?**

<!-- grade_id: cell-1742f85bb6e44a40, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ca7425aa985db6c1, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true(len(solution))
assert_equal(float(solution[0]),  40)
```

## Autó fogyasztás
Az autód 7 litert fogyaszt országúton és 9-et városban. A hévízi üdülésedre 170 kilómétert utazol országúton és 35-öt városban.

**Mennyit fogyaszt az autód csak oda?**

**Mennyit fogyaszt az autód oda-vissza?**

**Mennybe kerül a teljes út, ha 350 forint a benzin?**

<!-- grade_id: cell-7a92284df926610d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3a685528f6ed9482, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true(len(solution) == 3)
assert_equal(float(solution[0]),  15.05)
assert_equal(float(solution[1]),  30.1)
assert_equal(float(solution[2]),  10535.0)
```

# 2. feladatsor - változók

## Autó fogyasztás változókkal

**Oldd meg az "Autó fogyasztás" feladatot változókkal.** A szállás és étekezés a kétnapos hétvégére 35 000Ft. Összesen 70 000Ft-od van az utazásra.

**El tudsz-e menni az útra?**

<!-- grade_id: cell-09329f6eca268563, points: 0, grade: false, solution: true -->
```python
# A csak_oda, oda_vissza, teljes_ut_ara, el_tudok_e_menni nevu valtozoba szamoljuk ki az eredmenyt
csak_oda = 0
oda_vissza = 0
teljes_ut_ara = 0
el_tudok_e_menni = None

# YOUR CODE HERE
raise NotImplementedError()

print(csak_oda)
print(oda_vissza)
print(teljes_ut_ara)
print(el_tudok_e_menni)
```

<!-- grade_id: cell-6cb6b031ec2d8245, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal(csak_oda,  15.05)
assert_equal(oda_vissza,  30.1)
assert_equal(teljes_ut_ara,  10535.0)
assert_equal(el_tudok_e_menni,  True if el_tudok_e_menni is not None else None)
```

## Biztonságos jelszó

Ments el egy jelszót a "password" nevű változóban! A jelszónak legalább 8 karakter hosszúnak kell lennie, hogy biztonságosnak tekintsük. 

**Írj egy kifejezést, amely megvizsgálja, hogy biztonságos-e a megadott jelszó.**

<!-- grade_id: cell-64771b0dd0200ea6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-fac21acc0a8d7946, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true('password' in locals())
assert_true(len(solution))

if len(password) > 8:
    assert_equal(solution[0],  'True')
else:
    assert_equal(solution[0],  'False')    
```

## Kémia

1 dm<sup>3</sup> egységnyi alumínium tömege 2,7kg, a nikkelé pedig 8,8kg.

**Írasd ki: nehezebb-e 0,34 dm<sup>3</sup> nikkel, mint 1,23 dm<sup>3</sup> alumínium?**

_(Ne feledd, Pythonban tizedespontot kell használnod tizedesvessző helyett)_

<!-- grade_id: cell-c984a0b1b6ad2fcc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-27b9b29d44a50d50, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true(len(solution))
assert_equal(solution[0],  'True')
```

## Áram fogyasztás

A villanyórád 25 amperes. (Vagyis a 230V hálózati feszültséggel 25 * 230 Watt teljesítményt tud kiszolgálni)

A felkapcsolt világítást 5 db 30 wattos izzó biztosítja, a nyári meleg ellen bekapcsoltad a klímádat, ez jelenleg 1500 watton üzemel, illetve az 1200 wattos mosógéped is dolgozik.

**Írasd ki: Bekapcsolhatod-e a 2000 wattos vasalógépedet, vagy kell tőle tartanod, hogy lekapcsol a megszakító?**

<!-- grade_id: cell-2066d48477a99fc8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-207096d8c75c12a1, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true(len(solution))
assert_equal(solution[0],  'True')
```

# 3. feladatsor - String, operátorok, vezérlés

## Pultozás

**A felhasználó életkorát és italválasztását használva szolgáld őt ki!**

Kétféle italt ismerünk: sör és kóla
    
Ha a felhasználó 18 év alatti és sört kér, közöld vele hogy "Sajnos nem adhatok."

Ha a felhasználó 60 feletti és kólát kér, közöld vele, hogy "A koffein megemelheti a vérnyomását, sajnos nem adhatok."

Ha ismeretlen italt adott meg, akkor írd ki, hogy "Sajnos csak sört és kólát tudok adni."

Minden más esetben írd ki, hogy "Parancsoljon, itt a sör." vagy "Parancsoljon, itt a kóla".


_Ahol szükséges, ne felejtsd el konvertálni az értékeket_

<!-- grade_id: cell-0991a3ca87d6d839, points: 0, grade: false, solution: true -->
```python
# Ezeket atirva teszteld a megoldasodat
eletkor = 0  # Lehetseges ertekek: 0-999
ital = ""  # Lehetseges ertekek: "sör", "kóla"

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-fded93d3cf780074, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true(len(solution))

if eletkor < 18 and ital == "sör":
    assert_equal(solution[0], "Sajnos nem adhatok.")
elif eletkor > 60 and ital == "kóla":
    assert_equal(solution[0], "A koffein megemelheti a vérnyomását, sajnos nem adhatok.")
elif ital != "sör" or ital != "kóla":
    assert_equal(solution[0], "Sajnos csak sört és kólát tudok adni.")
else:
    assert_equal(solution[0], f"Parancsoljon, itt a {ital}.")
```

## Autó fogyasztás változókkal 2

Oldd meg a 2/a (üzemanyag fogyasztás) feladatot úgy, hogy nem előre megadott értékekkel (országúti fogyasztás, országúto megtett út, városi fogyasztás, városban megtetett út, benzin ára) dolgozol, hanem ezeket a felhasználótól kéred be.

_Ahol szükséges, ne felejtsd el konvertálni az értékeket_

<!-- grade_id: cell-0401409948bd40f0, points: 0, grade: false, solution: true -->
```python
# A kovetkezo valtozokba kerjuk be az adatokat
orszaguti_km = 0
orszaguti_fogyasztas = 0
varosi_km = 0
varosi_fogyasztas = 0
benzin_ara = 0

# A csak_oda, oda_vissza, teljes_ut_ara, el_tudok_e_menni nevu valtozoba szamoljuk ki az eredmenyt
csak_oda = 0
oda_vissza = 0
teljes_ut_ara = 0
el_tudok_e_menni = None

# YOUR CODE HERE
raise NotImplementedError()

print(csak_oda)
print(oda_vissza)
print(teljes_ut_ara)
print(el_tudok_e_menni)
```

<!-- grade_id: cell-a7de36810d308589, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal(csak_oda,  orszaguti_fogyasztas * orszaguti_km/100 + varosi_km/100 * varosi_fogyasztas)
assert_equal(oda_vissza,  csak_oda * 2)
assert_equal(teljes_ut_ara,  oda_vissza * 350)
assert_equal(el_tudok_e_menni,  70000 > teljes_ut_ara + 35000 if el_tudok_e_menni is not None else None)
```