# 1. feladatsor - Ciklusok

## Felhasználónév és jelszó
Írj programot, amely addig kér felhasználónevet és jelszót, amíg kellően erős adatokat nem adunk meg.

Felhasználónév legalább 6 karakter, jelszó legalább nyolc karakter, nem egyezik a felhasználónévvel.)

<!-- grade_id: cell-99da8cea42168c47, points: 0, grade: false, solution: true -->
```python
username = None
password = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-67adc014e0d56af4, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true(username is not None and password is not None)
assert_true(len(username) >= 6 and len(password) >= 8 and username != password)
```

## Sorminta

Készíts egy programot, ami képes a következő sormintát kirajzolni:

`sorok = 5`


`#
##
###
####
#####
`

<!-- grade_id: cell-7e8fc91bf1391d0c, points: 0, grade: false, solution: true -->
```python
sorok = 5

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8f82ba6dd106d32a, points: 1, grade: true, solution: false -->
```python
solution = %solution
assert_equal(solution[0], "#")
assert_equal(solution[1], "##")
assert_equal(solution[2], "###")
assert_equal(solution[3], "####")
assert_equal(solution[4], "#####")
```

## Szám kitalálás

Írj programot, amely 1 és 100 közt gondol egy véletlen számra. Tippet kér a felhasználótól, és kiírja, hogy
a gondolt szám nagyobb-e vagy kisebb. Ha a felhasználó eltalálja a gondolt számot, akkor a program gratulál, és a
futásának vége szakad (ezt a `Gratulalok eltalaltad` kiírásával tegye meg).

<!-- grade_id: cell-0058d833213a1604, points: 0, grade: false, solution: true -->
```python
import random

szam = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9ce09057eaf32821, points: 0, grade: true, solution: false -->
```python
solution = %solution

assert_true(type(szam) == int)
assert_true(len(solution))
assert_true(solution[-1] == 'Gratulalok eltalaltad')
```

## Számoló

Készíts egy programot, ami `while` ciklust használ, hogy a 1-1000 5el osztható számokat összeadja a `my_sum` változóba.

<!-- grade_id: cell-cc810269a1dd3747, points: 0, grade: false, solution: true -->
```python
my_sum = 0

# YOUR CODE HERE
raise NotImplementedError()
print(my_sum)
```

<!-- grade_id: cell-15c3ffa094c1ce7a, points: 0, grade: true, solution: false -->
```python
solution = %solution

assert_equal(my_sum, 100500)
```