# 1. feladatsor - funkciók

## Szökőév?
Készíts függvényt `is_szokoev` néven, amelyik adott évszámról eldönti, hogy az szökőév-e. Szökőév minden negyedik, nem szökőév minden századik, mégis az minden 400-adik. (2000-ben ezért volt szökőév.) A függvény visszatérési értéke legyen logikai típusú!

Tipp( % maradékos osztás operátor)

<!-- grade_id: cell-a3e8d85db2cd3829, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()


def main_szoko():
    print(is_szokoev(2005))
    print(is_szokoev(2000))
    print(is_szokoev(1980))
    print(is_szokoev(1900))

main_szoko()
```

<!-- grade_id: cell-d8bf68236e9f6fa6, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true('is_szokoev' in locals())
assert_false(is_szokoev(2005))
assert_true(is_szokoev(2000))
assert_true(is_szokoev(1980))
assert_false(is_szokoev(1900))
```

## Szökőév 2
Írj programot, amelyik a felhasználótól 4 évszámot kér, és mindegyikre kiírja, hogy szökőév-e! Használd az előbb megírt függvényt! Például:


? 2005

Nem szökőév.


? 2000

Szökőév.


? 1980

Szökőév.


? 1900

Nem szökőév.

<!-- grade_id: cell-84721ecb4b633c21, points: 0, grade: false, solution: true -->
```python
# Itt ha mar az elozo feladatban megirtuk az is_szokoev fuggvenyt, akkor azt tudjuk hivatkozni, nem kell idemasolni

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0bede469175bd15d, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true(len(solution) == 4)
for i in range(len(solution)):
    print(solution[i])
    assert_true(solution[i] in ["Szökőév.", "Nem szökőév."])
```

## „a” betűvel kezdődő sztring
Egy olyan függvényt kell írnod, amely paraméterként sztringek listáját kapja, legyen a neve `van_a_kezdetu`. Meg kell vizsgálnia a listában található sztringeket, és megválaszolnia ezt a kérdést:

Van-e olyan sztring, amelyik „a” betűvel kezdődik! Ennek is legyen logikai típusú a visszatérési értéke.


Ügyelj arra, hogy a függvény helyesen működjön abban az esetben is, ha üres sztring van a listában!

 Példák:

`["körte", "alma", "barack"]` – van „a” betűvel kezdődő szó.

`["dinnye", "papaja", "", "zeller"]` – nincs „a” betűvel kezdődő szó.

`[]` – nincs „a” betűvel kezdődő szó.


<!-- grade_id: cell-3d41946d709432c4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

def main_van_a():
    print(van_a_kezdetu([]))
    print(van_a_kezdetu(['atest', 'werwer', 'werwersd']))

main_van_a()
```

<!-- grade_id: cell-e3e4a0a8c8e2be2f, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_false(van_a_kezdetu([]))
assert_true(van_a_kezdetu(['atest', 'werwer', 'werwersd']))
assert_true(van_a_kezdetu(["körte", "alma", "barack"]))
assert_false(van_a_kezdetu(["dinnye", "papaja", "", "zeller"]))
```


# Madárnyelv
Adott az alábbi program, amely madárnyelven (mavadávárnyevelveven) írja ki a beírt szöveget.
```python
szoveg = input()
for c in szoveg:
    if c == 'a' or c == 'e' or c == 'i' or c == 'o' or c == 'u':
        print(c, 'v', c, sep="", end="")
    else:
        print(c, end="")
```

Írj függvényt (legyen a neve `maganhangzo`, amelyik megmondja egy betűről, hogy magánhangzó-e! Alakítsd át úgy a programot, hogy a megírt függvényt használod a main()-ben! A magánhangzós függvény megírható egyetlen egy sorból is (a fejlécén túl, természetesen).


Végezz el egy további átalakítást is: a madárnyelvesítés is legyen függvény (legyen a neve `madarnyelv`)! Vagyis írj egy olyan függvényt, amely paraméterként egy sztringet kap, visszatérési értéke is egy sztring; például madarnyelv("alma") == "avalmava". (A sztring összeállításához a += operátort tudod használni.) Használd fel ezt a függvényt a main()-ben az eredeti feladat megoldására!


Hogyan lehetne megoldani azt, hogy a nagybetűvel kezdődő szavakat is helyesen kezelje a program? Pl. az „Alma” szóra azt kell kiírnia, hogy „Avalmava”. Ehhez fel kell tudnia ismerni a nagybetűvel írt magánhangzókat is. Ugyanakkor az ismételt magánhangzó már a szó közepén van, tehát kisbetűsíteni kell.


<!-- grade_id: cell-f72580afa81807d9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

def main_madar():
    print(madarnyelv("Alma"))
    print(madarnyelv("Korte"))
    print(madarnyelv("Zeller"))
    print(madarnyelv("Macskafa"))

main_madar()
```

<!-- grade_id: cell-1e4aad7bd58acd1c, points: 1, grade: true, solution: false -->
```python
solution = %solution

for i in ['a', 'e', 'i', 'o', 'u']:
    assert_true(maganhangzo(i))
for i in ['b', 'c', 'd', 'f', 'g', 'h', 'j', 'k', 'l', 'm', 'n', 'p', 'q', 'r', 's', 't', 'v', 'w', 'x', 'y']:
    assert_false(maganhangzo(i))

assert_equal("Avalmava", madarnyelv("Alma"))
assert_equal("Kovorteve", madarnyelv("Korte"))
assert_equal("Zevellever", madarnyelv("Zeller"))
assert_equal("Mavacskavafava", madarnyelv("Macskafa"))
```

# Eltérő paraméterszámmal hívható függvény
A range() függvény három különféle paraméterezéssel hívható:

* `range(min, max)`: például a range(1, 5)-ből az 1, 2, 3, 4 számok potyognak ki.
* `range(min, max, step)`: adott lépésközzel, például range(0, 5, 2) paraméterek esetén 0, 2, 4.
* `range(max)`: egy paraméter esetén a kezdeti érték nulla lesz. Például range(3) esetén 0, 1, 2.

Írj függvényt `szamtani` néven, amely egy számtani sorozatot ír ki a képernyőre! A paraméterezése legyen az, mint a `range()` függvényé (figyeljunk arra, hogy `max` helyett a parameter neve legyen `_max`, stb), és pont ugyanazokat a számokat írd ki.

A függvény belsejében ne használd a range()-et! A feladat lényege az, hogy lásd, hogyan kell eltérő paraméterszámmal hívható függvényt írni.

<!-- grade_id: cell-5f9008984f9e2f8f, points: 0, grade: false, solution: true -->
```python
def szamtani(_min=None, _max=None, _step=None):
    # YOUR CODE HERE
    raise NotImplementedError()

def main_range():
    print(f'max range {szamtani(_max=10)}')
    print(f'min + max range {szamtani(_min=3, _max=10)}')
    print(f'min + max + step range {szamtani(_min=3, _max=10, _step=2)}')

main_range()
```

<!-- grade_id: cell-33a6c9043ac3ebd5, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], szamtani(_max=10))
assert_equal([3, 4, 5, 6, 7, 8, 9], szamtani(_min=3, _max=10))
assert_equal([3, 5, 7, 9], szamtani(_min=3, _max=10, _step=2))
```