# 1. feladatsor - Python alapok gyakorlása
## Vándorló madarak

Adott egy lista ami madár megfigyeléseket tartalmaz időrendi sorrendben. Egy bejegyzett szám az adott madár típusa.
`madarak = [2, 3, 2, 4, 5, 1, 2, 2, 3, 4, 6, 4, 2]`


Írj egy olyan Python programot ami összesíti egy `dict`-ben a madártípusokat, és hogy hányszor fordultak elő a mérés során.

<!-- grade_id: cell-84721ecb4b633c21, points: 0, grade: false, solution: true -->
```python
madarak = [2, 3, 2, 4, 5, 1, 2, 2, 3, 4, 6, 4, 2]
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0bede469175bd15d, points: 1, grade: true, solution: false -->
```python
solution = %solution
assert_equal(solution[0].strip(), "{2: 5, 3: 2, 4: 3, 5: 1, 1: 1, 6: 1}")
```

## CamelCase számoló
Készíts egy függvényt `camel_case_solver` néven, ami paraméternek egy karakterláncot kap, és megadja, hogy hány szó van összefűzve CamelCase módon a megadott stringben.

Pl:
`camel_case_solver("CamelCaseTestCase")` -> 4

<!-- grade_id: cell-1bc5a980cd583d68, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4d96a53f5f5e00fb, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_true('camel_case_solver' in locals())
assert_equal(camel_case_solver('CamelCaseTestCase'), 4)
assert_equal(camel_case_solver('alma'), 1)
assert_equal(camel_case_solver('almafa'), 1)
assert_equal(camel_case_solver('Alma'), 1)
assert_equal(camel_case_solver('AlmaFa'), 2)

```

## Doboz rajzoló

Készíts egy függvényt (`draw_box` néven), ami stringet készít, amellyel szöveget tudunk kiírni egy dobozba, pl:

```python
#########
# Hello #
#########
```

A függvénynek lehessen megadni:
* A doboz szélességét
* A doboz magasságát
* A kiírandó szöveget

A szöveget mindig helyezzük horizontális és vertikális középre. A függvény térjen vissza az eredménnyel, ne írjon ki.

Ha a megadott szélesség / magasságba nem férne be a szöveg, akkor térjen vissza `None` értékkel.

<!-- grade_id: cell-fa28ec479ae79a9b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
print(draw_box(15, 5, "hello"))
```

<!-- grade_id: cell-7b5025d3c59272e8, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal(draw_box(1, 1, "Hello"), None)
assert_equal(draw_box(5, 15, "Hello"), None)
assert_equal(draw_box(5, 7, "Hello"), None)
assert_equal(draw_box(15, 5, "hello"), "###############\n#             #\n#    hello    #\n#             #\n###############\n")
```