# 1. feladatsor - Fájl olvasás
## Hány betű?

Hozz létre egy szöveges fájlt `feladat.txt` néven, legyen a tartalma:

`Lorem ipsum dolor sit amet, consectetur adipiscing elit. Curabitur ac faucibus lacus, et commodo orci. Sed enim velit, malesuada gravida consequat vel, vulputate eget odio. Donec pellentesque luctus diam, quis pretium massa vestibulum quis. Sed a bibendum massa. Sed sollicitudin urna dignissim nulla rutrum, et venenatis est placerat. Vestibulum sollicitudin nisi eros, ac lacinia sem fermentum in. Nam imperdiet gravida facilisis. Phasellus lectus lectus, feugiat sit amet vulputate scelerisque, pretium non justo. Nulla vel ultrices arcu, sit amet consectetur purus. Nunc vel dolor sed est tempor venenatis eu eu risus. Nam lorem nulla, tincidunt id ex sit amet, feugiat tincidunt leo. Phasellus et felis eget arcu dignissim dignissim.`

Olvasd be a `feladat.txt` fájlt, és számold meg, hogy hány darab `a` betű van benne és írd ki az eredményt!

<!-- grade_id: cell-84721ecb4b633c21, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0bede469175bd15d, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal(int(solution[0].strip()), 45)
```

## Nagybetűsítés

Olvasd be a `feladat.txt` fájlt, és írd ki olyan módon, hogy minden szó legyen csupa nagybetű.


<!-- grade_id: cell-1bc5a980cd583d68, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4d96a53f5f5e00fb, points: 1, grade: true, solution: false -->
```python
solution = %solution
assert_equal(solution, ['LOREM IPSUM DOLOR SIT AMET, CONSECTETUR ADIPISCING ELIT. CURABITUR AC FAUCIBUS LACUS, ET COMMODO ORCI. SED ENIM VELIT, MALESUADA GRAVIDA CONSEQUAT VEL, VULPUTATE EGET ODIO. DONEC PELLENTESQUE LUCTUS DIAM, QUIS PRETIUM MASSA VESTIBULUM QUIS. SED A BIBENDUM MASSA. SED SOLLICITUDIN URNA DIGNISSIM NULLA RUTRUM, ET VENENATIS EST PLACERAT. VESTIBULUM SOLLICITUDIN NISI EROS, AC LACINIA SEM FERMENTUM IN. NAM IMPERDIET GRAVIDA FACILISIS. PHASELLUS LECTUS LECTUS, FEUGIAT SIT AMET VULPUTATE SCELERISQUE, PRETIUM NON JUSTO. NULLA VEL ULTRICES ARCU, SIT AMET CONSECTETUR PURUS. NUNC VEL DOLOR SED EST TEMPOR VENENATIS EU EU RISUS. NAM LOREM NULLA, TINCIDUNT ID EX SIT AMET, FEUGIAT TINCIDUNT LEO. PHASELLUS ET FELIS EGET ARCU DIGNISSIM DIGNISSIM.'])
```

# 2. feladatsor - Fájl írás
## Fájl írás

Írj ki 5 darab `x` karaktert az `x.txt` fájlba!

<!-- grade_id: cell-5f3ac85c68017818, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e4847263acee2df0, points: 1, grade: true, solution: false -->
```python
solution = %solution

import os

assert_true(os.path.isfile("x.txt"))
with open("x.txt", "r") as f:
    assert_equal("xxxxx", f.read())
```

## Fájlból fájlba

Olvasd be a `feladat.txt` fájlt, és írd ki a `tadalef.txt`-be úgy, hogy az egész fájl tartalmát megfordítod, pl:

`Lorem ipsum ...` -> `... muspi meroL`

<!-- grade_id: cell-e8113b35835eb69d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8b911d9d4fd99833, points: 1, grade: true, solution: false -->
```python
solution = %solution

with open("tadalef.txt", "r") as f:
    assert_equal(f.read(), ".missingid missingid ucra tege silef te sullesahP .oel tnudicnit taiguef ,tema tis xe di tnudicnit ,allun merol maN .susir ue ue sitanenev ropmet tse des rolod lev cnuN .surup rutetcesnoc tema tis ,ucra secirtlu lev alluN .otsuj non muiterp ,euqsirelecs etatupluv tema tis taiguef ,sutcel sutcel sullesahP .sisilicaf adivarg teidrepmi maN .ni mutnemref mes ainical ca ,sore isin niduticillos mulubitseV .tarecalp tse sitanenev te ,murtur allun missingid anru niduticillos deS .assam mudnebib a deS .siuq mulubitsev assam muiterp siuq ,maid sutcul euqsetnellep cenoD .oido tege etatupluv ,lev tauqesnoc adivarg adauselam ,tilev mine deS .icro odommoc te ,sucal subicuaf ca rutibaruC .tile gnicsipida rutetcesnoc ,tema tis rolod muspi meroL")
```

## Táblázatos 1

Hozz létre szövegszerkesztővel egy `tablazatos.txt` fájlt.
Írd bele a következőket:

```csv
Termékek|Mennyiség|Egységár|Összérték
alma|50|300|15000
banán|40|400|16000
citrom|30|500|15000
dió|20|2500|50000
```

Írj egy Python programot ami ebből a fájlból beolvassa az adatokat. Tárold az adatokat egy listában. Minden sortból keletkezzen egy egy lista. Az elválasztó karakter a `|` karkter legyen. Figyelj arra, hogy a whitespace karaktereket (pl `\n`) ne tárold el.

Írd ki az igy keletkezett listát.


<!-- grade_id: cell-6b69c0aa2e994a75, points: 0, grade: false, solution: true -->
```python
termekek = []

# YOUR CODE HERE
raise NotImplementedError()

print(termekek)
```

<!-- grade_id: cell-9825b1a3e6c91d00, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal(solution[0], "[['alma', 50, 300, 15000], ['banán', 40, 400, 16000], ['citrom', 30, 500, 15000], ['dió', 20, 2500, 50000]]")
```

## Táblázatos 2

Az előző feladatban keletkezett sorok listájából készíts egy olyan változatot amiben megduplázol minden értéket.

Az így keletkezett lista tartalmát írd ki egy `dupla_tablazatos.txt` fájlba.

<!-- grade_id: cell-9086cad1053b7716, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d6a854948b5a189e, points: 1, grade: true, solution: false -->
```python
solution = %solution

import os

assert_true(os.path.isfile("dupla_tablazatos.txt"))
with open("dupla_tablazatos.txt", "r") as f:
    idx = 0
    for line in f:
        name = line.split("|")[0]
        assert_equal(name, termekek[idx][0] * 2)
        idx += 1
```