# 1. feladatsor - Lista származtatás (comprehension)

## For ciklus átírás

Írd át a következő for ciklust list comprehension formába.

```
lista = [1, 1, 2, 3, 5, 8, 12]
uj_lista = []
for i in lista:
    uj_lista.append(i ** 2)
print(uj_lista)
```

<!-- grade_id: cell-84721ecb4b633c21, points: 0, grade: false, solution: true -->
```python
lista = [1, 1, 2, 3, 5, 8, 12]
# YOUR CODE HERE
raise NotImplementedError()
print(uj_lista)
```

<!-- grade_id: cell-0bede469175bd15d, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal(solution[0].strip(), "[1, 1, 4, 9, 25, 64, 144]")
```

## Válogatás

Adott egy lista:
`lista = ["alma", "korte", "zsiraf", "akarat", "haz", "alaptetel"]`

Készíts ebből list comprehensionnel egy új listát amiben csak az `a` betűs szavak vannak benne, és írd ki.


<!-- grade_id: cell-1bc5a980cd583d68, points: 0, grade: false, solution: true -->
```python
lista = ["alma", "korte", "zsiraf", "akarat", "haz", "alaptetel"]
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4d96a53f5f5e00fb, points: 1, grade: true, solution: false -->
```python
solution = %solution
assert_equal(solution[0].strip(), "['alma', 'akarat', 'alaptetel']")
```

## Mit vegyek

Adott egy bevásárlólista:
```python
bevas = {
    "alma" : 12,
    "korte" : 66,
    "barack" : 14,
    "ragogumi": 2
}
```

Készíts egy listát list comprehension használatával a beváslistán szereplő tételek neveiből amik értéke nagyobb mint `13` és írd ki ezeket.

<!-- grade_id: cell-b6c1d03bdec0e894, points: 0, grade: false, solution: true -->
```python
bevas = {
    "alma" : 12,
    "korte" : 66,
    "barack" : 14,
    "ragogumi": 2
}

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e5f400fa0276104f, points: 1, grade: true, solution: false -->
```python
solution = %solution
assert_equal(solution[0].strip(), "['korte', 'barack']")
```