# 1. feladatsor - Reguláris kifejezések
## Szöveget keresünk

Adottak a következő patternek:

`patterns = [ 'fox', 'dog', 'horse' ]`

És a következő szöveg:

`text = 'The quick brown fox jumps over the lazy dog.'`


Írj egy programot, ami egyesével megpróbálja megkeresni a patterneket a szövegben, és mindegyiknél kiírja, hogy `Benn van` vagy `Nincs benne`.

<!-- grade_id: cell-84721ecb4b633c21, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0bede469175bd15d, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal(solution[0].strip(), "Benne van")
assert_equal(solution[1].strip(), "Benne van")
assert_equal(solution[2].strip(), "Nincs benne")
```

## URL keresés

Adott egy szöveg amiben url-ek is vannak:

`text = "Látogass el a https://valami.hu oldalra, vagy nézd meg a legújabb ajánlatunkat a http://masikvalami.hu oldalon"`

Írj egy programot, ami kiválogatja a linkeket a szövegből és kiírja ezeket.



<!-- grade_id: cell-1bc5a980cd583d68, points: 0, grade: false, solution: true -->
```python
text = "Látogass el a https://valami.hu oldalra, vagy nézd meg a legújabb ajánlatunkat a http://masikvalami.hu oldalon"

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4d96a53f5f5e00fb, points: 1, grade: true, solution: false -->
```python
solution = %solution
assert_equal(solution[0], "['https://valami.hu', 'http://masikvalami.hu']")
```

# Szöveg cenzúra

Adott egy szöveg:

`text = "jó feketeség volt előttünk, fekete napra ébredtünk, mély fekete víz tükrét néztük"`

A cenzornak nem tetszik a sok `fekete` szó bármilyen formája, írj egy programot ami reguláris kifejezések használatával kitörli ezeket és kiírja az új szöveget.


<!-- grade_id: cell-5f3ac85c68017818, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e4847263acee2df0, points: 1, grade: true, solution: false -->
```python
solution = %solution

assert_equal(solution[0], "jó volt előttünk, napra ébredtünk, mély víz tükrét néztük")
```