# 1. feladatsor - Dátum kezelés
## Visszaszámláló
Írj egy Python programot ami egy visszaszámláló órát valósít meg, nevezd el a függvényt `countdown` néven.

Válassz egy eseményt ami tudod mikor van, mondjuk újév.

Írd ki, hogy mennyi év, hó, nap, perc, másodperc van még hátra az eseményig, pl: `200 days, 2:31:03.938311`

<!-- grade_id: cell-84721ecb4b633c21, points: 0, grade: false, solution: true -->
```python
import time
from datetime import datetime

# a nevezetes datum
new_year_2022 = datetime(year = 2022, month = 1, day = 1, hour = 0, minute = 0, second = 1)

# YOUR CODE HERE
raise NotImplementedError()
countdown(3)
```

<!-- grade_id: cell-0bede469175bd15d, points: 1, grade: true, solution: false -->
```python
import re
from datetime import timedelta
solution = %solution

def parse(s):
    if 'day' in s:
        m = re.match(r'(?P<days>[-\d]+) day[s]*, (?P<hours>\d+):(?P<minutes>\d+):(?P<seconds>\d[\.\d+]*)', s)
    else:
        m = re.match(r'(?P<hours>\d+):(?P<minutes>\d+):(?P<seconds>\d[\.\d+]*)', s)
    return timedelta(**{key: float(val) for key, val in m.groupdict().items()})

countdown_result = parse(solution[0].strip())
delta = new_year_2022 - datetime.now()
assert_true(countdown_result - delta < timedelta(minutes=1))

```